---
title: 'Initial Course'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`
Complete the code the print "Hello World"

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("Hello World")
```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 94383eab17
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
